# Data Structures and Processing

## Week3: NumPy, and Matplotlib

### Task 1: Custom NumPy Functions

Write down a function `heaviside` which takes a number (`int` or `float`) and returns 1 if it is positive, 0 otherwise.

<!-- grade_id: cell-5cb06a6dd536e09a, points: 0, grade: false, solution: true -->
```python
# Your solution goes here.
def heaviside(num):
    
    if num > 0:
        
        return 1
    
    else:
        
        return 0

```

<!-- grade_id: cell-1aaea127bbe9a155, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert heaviside(0.1) == 1
assert heaviside(-0.1) == 0
assert heaviside(0) == 0
assert heaviside(4) == 1
### END HIDDEN TESTS
```

At this point, your function `heaviside` will return a `ValueError` if a NumPy array is provided as an input to it.  We would like to extend this function so that it handles any `ndarray` whose dtype is `int` or `float`.

Extend the `heaviside` function to handle the `ndarray`.

<!-- grade_id: cell-b1fe8aefbdc9898f, points: 0, grade: false, solution: true -->
```python
import numpy as np

def heaviside(num):
    
    if isinstance(num, (int, float)):
        
        if num > 0:
            
            return 1
        
        else:
            
            return 0
        
    elif isinstance(num, np.ndarray) and np.issubdtype(num.dtype, np.number):
        
        return np.where(num > 0, 1, 0)
    
    else:
        
        raise ValueError("Invalid Input!")
```

<!-- grade_id: cell-c9085a23523dcd7c, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import numpy as np
assert np.array_equal(heaviside(np.array([-1.2, 0.2,3])), np.array([0, 1, 1]))
```

### Task 2: Aggregate Functions

There are a lot of aggregate functions available in NumPy. Some of the variable useful ones include: `np.sum`, `np.mean`, `np.std`, `np.max`, `np.min` and so on.  For a detailed explanation of the parameters these function take in, have a look at their documentation using `help`.

Write down a function `mean_squared_values` that takes an `ndarray` as an argument and returns the mean value of the squares of elements in it.

<!-- grade_id: cell-e6aa4fb4daec46cc, points: 0, grade: false, solution: true -->
```python
import numpy as np

def mean_squared_values(arr):
    
    sqrt = np.square(arr)
    
    meanres = np.mean(sqrt)
    
    return meanres
```

<!-- grade_id: cell-ed3ec8b79ec4fb45, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import numpy as np
x_utest = np.random.normal(size=(2,3))
assert mean_squared_values(x_utest) == np.mean(x_utest**2)
### END HIDDEN TESTS
```

### Task 3: I/O and Visualization

In this task, we plan to read a data to an `ndarray`, process it, visualize the processed data, and write down to a file.

#### Step 1: Reading Data From a File

Consider the file `data.txt` in the current directory, and read its content to an array named `arr`.

Write down a function `read_values` that takes a `CSV` file `filename` as an input and an `index` as a second argument, and returns the `index` column from the `CSV` file as a NumPy array.

<!-- grade_id: cell-3590743b79d3ee18, points: 0, grade: false, solution: true -->
```python

    
```

<!-- grade_id: cell-9c00adcc35e85653, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import numpy as np
assert read_values("data.txt") == np.loadtxt("data.txt", delimiter=",", usecols=0)
### END HIDDEN TESTS
```

Use your function to read in values to two variables `x_values` and `y_values`.

<!-- grade_id: cell-8845595b70b6b28d, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-210e1768fb0c3ffd, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert x_values == np.loadtxt("data.txt", delimiter=",", usecols=0)
assert y_values == np.loadtxt("data.txt", delimiter=",", usecols=1)
### END HIDDEN TESTS
```

### Step 2: Processing

At this point, we are choosing a simple only to demonstrate the processing before plotting.  The data read in does not reveal something useful because it is randomly generated.

The reader should see that any type of processing is possible at this step.

Let us say that we do not want to change the variable `x_values`, but we do want to process the `y_values` by applying an exponential function `np.exp` from the NumPy library.

<!-- grade_id: cell-c8686dbb972435e7, points: 0, grade: false, solution: true -->
```python
# Your solution goes here.  Replace `None` below with appropriate 

y_values = None
```

<!-- grade_id: cell-4932dc90b5251480, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert y_values == np.exp(np.loadtxt("data.txt", delimiter=",", usecols=1))
### END HIDDEN TESTS
```

#### Step 3: Visualization

Take the `x_values` and `y_values` and plot them using `plt.plot`.  Given extra directive of `"bo"` to the plot command.

<!-- grade_id: cell-63928df5f03f64e3, points: 0, grade: false, solution: true -->
```python
# Your solution to plot goes here

```

#### Step 4: Writing Data to File

Write down a function `data_to_csv` that takes three arguments: `x_values`, `y_values` and `filename`, and writes down the arrays as a CSV file to `filename`.  The delimiter is `","`.

<!-- grade_id: cell-001fb10d6ae56bd0, points: 0, grade: false, solution: true -->
```python
# Your solution to 
def data_to_csv(x_values,y_values, filename):
    pass
```

<!-- grade_id: cell-8a247410e7de9584, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert 2 == 1+1
### END HIDDEN TESTS
```